<a href="https://colab.research.google.com/github/anythingispossible19/SP19/blob/master/debuggin_question_10_exercise_sheet_1exceptions_for_resolution.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Exercise Sheet 1
## Instructions
This exercise sheet counts towards your overall grade for the unit. Please refer to Engage for the most up-to-date administrative details on grading and deadlines.

Each exercise sheet will be made up of several programming challenges in Python. The content of the exercises will primarily focus on the material from the relevant week, but concepts from previous weeks will also be required, and you may even need to pick up new concepts that are introduced within the exercise itself – this is a normal part of programming.

Each exercise is spread across three Jupyter notebook cells:
* A text cell (like this one) introducing the exercise
* A code cell where you can fill your solution, which may include some code already pre-filled for you to modify
* A code cell containing tests

Each test cell will normally include visible tests that you must run yourself to check your submission is working and in the format we expect. If you can run the cell and *nothing* happens, the tests are passing. These tests are mandatory, and do not typically include challenging test cases. If your submission does not pass these tests then it may receive zero marks for the exercise. 

The test cell also contains hidden tests, and these will be used to actually determine your grade for the exercise. So you **must** also test your code on additional inputs that you determine yourself. This is part of writing code, you need to check that it works. 

Unless otherwise specified, you *do not* need to ensure that your code is robust for unusual inputs. If the exercise states that the input will be a number, you can expect the input to be a number, you do not need to check if it is a string.

You can test your code in the middle code cell or add additional cells to the notebook, but *please delete these before your final submission*, which should just be the exercise code itself.

In general, unless the exercise instructions state otherwise, you can feel free to modify any skeleton code that has been provided *provided* the test cell still runs successfully.

## Questions
### Question 1
Write a function that takes a single number, adds 3 to that number, then multiplies the result by 2, then subtracts 4, then subtracts twice the original number, then adds 3, and returns the result.

<!-- grade_id: cell-8ec771280bdacf20, points: 0, grade: false, solution: true -->
```python
def process_number(x):
    
    
    raise NotImplementedError()
```

<!-- grade_id: cell-37523cf9c5e087d0, points: 10, grade: true, solution: false -->
```python
# This is a test cell, if it runs without errors then your code is in the correct format for submission
# but you must still perform your own testing to ensure you are happy your answer is correct.
assert(process_number(5) == 5)
```

### Question 2
Write a function which converts a length given in centimetres to feet. You will likely need to look up the formula for this conversion online, try to be as exact as possible!

<!-- grade_id: cell-62ec441834039e8b, points: 0, grade: false, solution: true -->
```python
def centimetres_to_feet(cm):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8d322b244bce8acb, points: 10, grade: true, solution: false -->
```python
assert(centimetres_to_feet(0) == 0)
```

### Question 3
Write a function that takes two integers `x` and `y`. It should return a string containing the integer `x` repeated `y` times. See the provided tests below for examples.

<!-- grade_id: cell-a80faf28af2fff25, points: 0, grade: false, solution: true -->
```python
def repeat_int_str(x, y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-abad8e6be2aa33b3, points: 10, grade: true, solution: false -->
```python
assert(repeat_int_str(0, 4) == "0000")
assert(repeat_int_str(-5, 3) == "-5-5-5")
```

### Question 4
Now write a function similar to the one above, that takes two integers `x` and `y`, but this time produces a string *of length* `y` made up of repeated copies of `x`.

<!-- grade_id: cell-9ce26ef0a58e99bb, points: 0, grade: false, solution: true -->
```python
def repeat_int_length(x, y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b0c00ceb2ed4ab1b, points: 10, grade: true, solution: false -->
```python
assert(repeat_int_length(10, 4) == "1010")
assert(repeat_int_length(-5, 3) == "-5-")
```

### Question 5
One more time, write another function which takes arguments `x` and `y`, and produces a string of length `y` made of repeated copies of `x`, ***but*** this time, if `x` is negative do not repeat the negative sign, only the digits. Again, see the examples in the tests below for clarification.

<!-- grade_id: cell-ece9e5b5e185e772, points: 0, grade: false, solution: true -->
```python
def repeat_int_no_minus(x, y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7ba13193fa2aeee2, points: 10, grade: true, solution: false -->
```python
assert(repeat_int_no_minus(-5, 5) == "-5555")
assert(repeat_int_no_minus(-343, 3) == "-34")
```

### Question 6
Remember that strings can be compared using `<` and `>`. If you are comparing characters that are the same case (both UPPER CASE or both lower case), then `<` and `>` will compare the characters alphabetically. So `"a" < "z"` evaluates to `True`. 

Write a function that takes an input string all in lower case and finds the smallest (alphabetically first) *letter* from (a-z). Importantly, note that some special characters (like spaces) might be considered “smaller” and so must be ignored, i.e. `" " < "a"` evaluates to `True`. If the input string contains no lower case letters from a-z, return an empty string.

<!-- grade_id: cell-62b319d167ceab6f, points: 0, grade: false, solution: true -->
```python
def smallest_letter(string):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-48d46b3c2804273d, points: 10, grade: true, solution: false -->
```python
assert(smallest_letter("hello world") == "d")
assert(smallest_letter(" ") == "")
```

### Question 7
The reason why we can compare strings alphabetically using `<` and `>`, but only for characters of the same case, is because it is *actually* comparing the [character values](https://en.wikipedia.org/wiki/List_of_Unicode_characters#Latin_script). Python uses Unicode for its characters, but you might also hear people refer to ASCII character values too, since Unicode is backwards compatible and uses the same values for basic Latin script, numerals, and the symbols you find on a UK keyboard.

You can convert a single-character string into its character value using the inbuilt function `ord`, and you can convert back from a number into a single-character string using `chr`. So `ord("A") == 65` and `chr(65) == "A"`. Since each character's value is sequential (A is 65, B is 66, …), if you think about it, you can use these functions to convert between cases. You can also use string methods like `.lower()` and `.upper()`.

Now, armed with this new information, write a function which works the same as the previous question except it produces the smallest alphabetical letter *of either case*. If the string contains an upper and lower case character which is alphabetically first, return the value of the character which occurs first in the string. The case of the return value must match the input. See the tests for clarifying examples.

<!-- grade_id: cell-b1729d6895302bf6, points: 0, grade: false, solution: true -->
```python
def smallest_ignore_case(string):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-96eff72adde04801, points: 10, grade: true, solution: false -->
```python
assert(smallest_ignore_case("Hello World") == "d")
assert(smallest_ignore_case("HeLLo WorLD") == "D")
assert(smallest_ignore_case("What an AARDVARK!") == "a")
```

### Question 8
Write a function which takes a string as input and returns the biggest distance between any two letters (a-z ignoring case). Distance in this case means the number of letters apart in the alphabet, so the distance between `"a"` and `"b"` is `1`, the distance between `"a"` and `"Z"` is `25`. If the string contains no letters return `-1`.

<!-- grade_id: cell-cdb03a89c16be097, points: 0, grade: false, solution: true -->
```python
def biggest_letter_distance(string):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4e3bc1006f3cde9a, points: 10, grade: true, solution: false -->
```python
assert(biggest_letter_distance("Hello World") == 19)
assert(biggest_letter_distance("!!") == -1)
```

### Question 9
In the following cell the programmer has been experimenting with writing a function that can take either two or three arguments. However, they have made a mistake with their syntax. Find the mistake and fix the code so that it behaves the way they are expecting.

<!-- grade_id: cell-8b1227a02347648d, points: 0, grade: false, solution: true -->
```python
def my_function(a, b):
    if a < b:
        return a ** 2 + b
    else:
        return b

def my_function(a, b, c):
    if c < a and c < b:
        return c ** 3 + a ** 2 + b
    else:
        return my_function(a, b)
```

<!-- grade_id: cell-cf0fd922b0a9b735, points: 10, grade: true, solution: false -->
```python
assert(my_function(4, 5, 3) == 48)
assert(my_function(2, 1) == 1)
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-14-f702f9c734d7> in <module>
      1 assert(my_function(4, 5, 3) == 48)
----> 2 assert(my_function(2, 1) == 1)

TypeError: my_function() missing 1 required positional argument: 'c'
```

### Question 10
In the following code cell, a programmer has tried to re-implement the `*` and `//` operators as recursive functions that only use addition and subtraction; no other arithmetic operators (obviously `*` and `//` are not allowed) and no functions like `abs`. We saw an example of multiplication using addition in the unit material already using a loop, and the recursive version is a similar idea. Both functions should work for all reasonably-sized integer inputs `x` and `y`, the programmer is not trying to make them efficient. They should produce *exactly* the same results as the `*` and `//` operators.

Unfortunately the programmer has made some errors in their code. For this exercise, debug the code to find and fix all of the errors. Bear in mind that just getting the code to run does not mean you are done, as usual you will need to test the code with various inputs yourself before submitting. This will likely require tracing the code by hand keeping track of the values, even through layers of recursion!

<!-- grade_id: cell-09b65380a7494ba7, points: 0, grade: false, solution: true -->
```python
def multiply(x, y):
    if y != 0:
        return x + multiply(x, y - 1)
    else:
        return 0
    
def divide(x, y):
    print("")
    if x < 0 and y < 0:
        print("divide function, 'x < 0 and y < 0:'")
        print("x: ", x)
        print("y: ", y)
        print("")
        return divide(0 - x, 0 - y)
    if x < 0 or y < 0:
        print("divide function, 'x < 0 or y < 0:'")
        print("x: ", x)
        print("y: ", y)
        print("")
        return 0 - 1 - divide(0 - x, y)
    else:
        print("divide, 'else:'")
        print("x: ", x)
        print("y: ", y)
        print("")
        return divide_helper(x, y, 0)

def divide_helper(x, y, temp):
    if y >= x or y <= x and y < 0: # Need to work on this line 
    # if y >= x:
        print("divide_helper, 'y >= x:''")
        print("x: ", x)
        print("y: ", y)
        print("answer: ", temp)
        print("remainder: ", x)
        print("")
        return temp
    else:
        print("divide_helper, 'else:''")
        print("x: ", x)
        print("y: ", y)
        print("temp: ", temp)
        print("")
        return divide_helper(x - y, y, temp + 1)
```

<!-- grade_id: cell-e1ddc5f3ee96a7d5, points: 10, grade: true, solution: false -->
```python
assert(multiply(3, 5) == 3 * 5)
assert(divide(13, 3) == 13 // 3)
# note: using * or // (or similar) in your submission will result in zero marks!
```

Output:
```
divide, 'else:'
x:  13
y:  3

divide_helper, 'else:''
x:  13
y:  3
temp:  0

divide_helper, 'else:''
x:  10
y:  3
temp:  1

divide_helper, 'else:''
x:  7
y:  3
temp:  2

divide_helper, 'else:''
x:  4
y:  3
temp:  3

divide_helper, 'y >= x:''
x:  1
y:  3
answer:  4
remainder:  1
```

```python
assert(divide(25, 5) == 25 // 5)
assert(divide(-25, -5) == -25 // -5)
assert(divide(25, -5) == 25 // -5)
assert(divide(-25, 5) == -25 // 5)
assert(divide(26, 5) == 26 // 5)
assert(divide(-26, -5) == -26 // -5)
assert(divide(26, -5) == 26 // -5)
assert(divide(-26, 5) == -26 // 5)
```

Output:
```
divide, 'else:'
x:  25
y:  5

divide_helper, 'else:''
x:  25
y:  5
temp:  0

divide_helper, 'else:''
x:  20
y:  5
temp:  1

divide_helper, 'else:''
x:  15
y:  5
temp:  2

divide_helper, 'else:''
x:  10
y:  5
temp:  3

divide_helper, 'y >= x:''
x:  5
y:  5
answer:  4
remainder:  5

AssertionError                            Traceback (most recent call last)
<ipython-input-26-977bfad3a6ac> in <module>()
----> 1 assert(divide(25, 5) == 25 // 5)
      2 assert(divide(-25, -5) == -25 // -5)
      3 assert(divide(25, -5) == 25 // -5)
      4 assert(divide(-25, 5) == -25 // 5)
      5 assert(divide(26, 5) == 26 // 5)

AssertionError:
```

```python

```