<a href="https://colab.research.google.com/github/anythingispossible19/SP19/blob/master/exercise_sheet_3New_Copy1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Exercise Sheet 3
## Instructions
The instructions for this exercise sheet are the same as the previous, so this section is just a summary to serve as a reminder. Refer to the material on Engage and at the start of Exercise Sheet 1 for more detail, and if you are unsure ask on the Q&A forum.

This exercise sheet counts towards your overall grade for the unit.

Complete each question in the code cell provided, which will usually contain some skeleton code to help you start the question. Unless specified otherwise you may change the skeleton code as you wish, but your code must pass the formatting tests in the following code cell – if the cell runs without errors then your code is eligible for submission. However these will only test the basics of the question. Your code will be subject to additional hidden tests which will determine your grade. Make sure you thoroughly test your own code to check you think it works correctly before submission.

**Note**: some of these exercises require manipulation of files. Provided you only modify the files as the exercises instruct you to, you should not have any compatibility problems. However errors might occur if you make a mistake and edit a file that we are only expecting you to read. So you may wish to test your code works with a fresh copy of the default files from Engage. You can entirely replace the contents of the folder called `exercise-files` with a fresh copy. You still only need to submit this `.ipynb` file when you are done, no extra files.

## Questions
### Question 1
In the UK, the National Lottery requires you to pick 6 numbers from 1 to 59 inclusive. There is a well known formula for calculating how many ways you can draw $x$ values from $y$ possible values, these are called [combinations](https://en.wikipedia.org/wiki/Combination). 

For this exercise, write a function which takes two positive integers `x` and `y`, where $x\le y$, and returns the probability of winning the lottery jackpot, for `x` balls which are drawn from values `1-y` inclusive. The probability is equal to 1 divided by the number of possible combinations for that many balls.

*Rather than calculate this manually, you are encouraged to use the* [`math`](https://docs.python.org/3/library/math.html) *module.*

<!-- grade_id: cell-c58008015c7f239d, points: 0, grade: false, solution: true -->
```python
from math import comb

def lottery_probability(x, y):
    
    #math.comb(n, k)
    #Return the number of ways to choose k items from n items without repetition and without order.
    #Evaluates to n! / (k! * (n - k)!) when k <= n and evaluates to zero when k > n.
    #https://docs.python.org/3/library/math.html  
    combinations = comb(y, x)
    
    # return probability of a winning selection 
    return 1 / combinations

```

<!-- grade_id: cell-0f04eb504a1079d4, points: 20, grade: true, solution: false -->
```python
assert(lottery_probability(6, 59) == 1/45057474)
```

### Question 2
Write a function that takes two integers `x` and `y` and returns a list containing every possible *combination* of lottery draw as a tuple, where the lottery draws `x` balls with the values `1-y` inclusive.

So, if you set `x=6` and `y=59`, you should return a list of length 45057474, where each item is a different possible lottery draw, e.g. `(1, 2, 3, 4, 5, 6)` is one such draw. Order does not matter, so `(2, 3, 4, 5, 6, 1)` is considered the same draw.

*Rather than calculate this manually, you are encouraged to use the* [`itertools`](https://docs.python.org/3/library/itertools.html) *module.*

<!-- grade_id: cell-dddb8c4b6c7f9884, points: 0, grade: false, solution: true -->
```python
from itertools import combinations 

def all_lottery_draws(x, y):    
   
    rangeData = range(1, y + 1)
        
    #combinations(iterable, r)
    #Return r length subsequences of elements from the input iterable.
    #https://docs.python.org/3/library/itertools.html
    combinationData = list(combinations(rangeData, x))    
    return combinationData
```

<!-- grade_id: cell-a86a034f0371c48b, points: 20, grade: true, solution: false -->
```python
assert(all_lottery_draws(2, 3) == [(1, 2), (1, 3), (2, 3)])
```

### Question 3
In the code below, the programmer has tried to write a function which accepts a two-dimensional list representing a mathematical matrix and a number, where the purpose of the function is to return a new list with each element multiplied by that number.

Unfortunately, the code does not work as you would expect. First of all, run the function – you can use the given test case to work out why the results are wrong. Then, find and fix the error.

Feel free to search online to find clues for why this is not working, although you have seen an explanation of the mechanics in the unit material. You may wish to use an additional module in your solution to fix the error, and you are in fact encouraged to do so, but you must work out what the module is for yourself!

<!-- grade_id: cell-e99128cf596349bc, points: 0, grade: false, solution: true -->
```python
import copy

def matrix_scale(matrix, x):    
    # deep copy needed, as matrix variable is updated using shallow matrix.copy() function 
    new_matrix = copy.deepcopy(matrix)   
    
    for i in range(len(matrix)):       
        for j in range(len(matrix[i])):           
            new_matrix[i][j] = new_matrix[i][j] * x
           
    return new_matrix
```

<!-- grade_id: cell-03851bb5a2450311, points: 20, grade: true, solution: false -->
```python
my_matrix = [[1, 0, 0],
             [0, 1, 0],
             [0, 0, 1]]

double = matrix_scale(my_matrix, 2)
negation = matrix_scale(my_matrix, -1)

assert(double == [[2, 0, 0], [0, 2, 0], [0, 0, 2]])
assert(negation == [[-1, 0, 0], [0, -1, 0], [0, 0, -1]])

```

### Question 4
Write a function which takes two strings, which are both filenames for existing text files. Your function should compare the two files' contents, and return a list containing every line number where the files are different. Start numbering the files from 1 (the first line is line 1). If one file is longer than the other, every line not present should be considered unequal. If the files are identical, return an empty list. 

Note: in the unit material, we showed a file which did not end in a linefeed character on its final line, but this is generally considered bad form, all text files should be terminated by a linefeed. All inputs to your function *will* terminate in a linefeed character, but either way you do not need to worry about an "empty line" at the end of the file.

<!-- grade_id: cell-79f41a862640b444, points: 0, grade: false, solution: true -->
```python
def diff(filename1, filename2):
        
    master_file = open(filename1).read()
    compare_file = open(filename2).read()

    lines_master = master_file.splitlines()       
    lines_compare = compare_file.splitlines()
    
    different_lines = [] #A list type here is used to collect the line numbers  
    
    #Python enumerate() function can be used to iterate the list in an optimized manner. 
    #The enumerate() function adds a counter to the list or any other iterable and returns 
    #it as an enumerate object by the function. 
    #Thus, it reduces the overhead of keeping a count of the elements while the iteration operation.
    
    for i, line in enumerate(lines_master):
       if not line in lines_compare:
          different_lines.append(i + 1) #Manage the indexing of 0 by adding 1

    return(different_lines)
```

<!-- grade_id: cell-4fdac4989a8a30c2, points: 20, grade: true, solution: false -->
```python
assert(diff('./exercise-files/text_a.txt', './exercise-files/text_b.txt') == [3, 4, 6])
assert(diff('./exercise-files/text_a.txt', './exercise-files/text_a.txt') == [])
```

### Question 5
In this unit, we use a Jupyter extension called nbgrader to write these assignments. This adds extra metadata into each cell of the notebook to enable us to easily run tests on your code. Each cell that nbgrader knows about is given a cell ID to uniquely distinguish it – this helps identify which cell is which if, for example, you add additional cells into the notebook.

Each Jupyter notebook is stored in a format called [JSON](https://en.wikipedia.org/wiki/JSON). This is a structured way of writing data which is easily machine-readable, and somewhat human-readable too. 

While I do not recommend editing the JSON encoded version of any of your assignments (if you break the tests you may get zero marks), you can read and decode the files using the Python [`json`](https://docs.python.org/3/library/json.html) module.

In the cell below, write a function which takes a filename as a string. If the file is a Jupyter notebook that has nbgrader metadata in the first cell, then return the ID of that cell. Otherwise, return an empty string.

<!-- grade_id: cell-717b02ca63f2e9c0, points: 0, grade: false, solution: true -->
```python
import json

def get_first_cell_id(filename):
    
    #Is the file a jupyter notebook? 
    if not filename.endswith('.ipynb'):
        #If not a jupyter notebook, return empty string, as required by test case  
        return ("")
    else:
        #If so, open the JSON 
        with open(filename) as file:
            notebookData = json.load(file)
        
        #Raw data
        #print(notebookData)
        
        #Check the 1st cell metadata for nbgrader; 0 indexedfor 1st 
        nbgraderFlag = True if "nbgrader" in notebookData['cells'][0]['metadata'] else False

        #If nbgrader is there, get the grade_id value, which should be cell-06eb0ab990e028e8 
        #(after visual inspection in Notepad++, and test case inspection)
        if nbgraderFlag == True:
            #Set the grade_id value to target 
            target = notebookData['cells'][0]['metadata']['nbgrader']['grade_id']
            #return target, as required by test case              
            return (target)
        else:
            #JSON yet no nbgrader present 
            return ("")
```

<!-- grade_id: cell-25e4fda9a2fc80f4, points: 20, grade: true, solution: false -->
```python
assert(get_first_cell_id("exercise-sheet-3.ipynb") == 'cell-06eb0ab990e028e8')
assert(get_first_cell_id("./exercise-files/text_a.txt") == '')
```

```python

```